#   `lab04`—Expressions, Data, & Formulae

**Objectives**

-   Access parts of a data value, like a string, using indexing.  <!-- dimdak -->
-   Compose expressions and statements using Python syntax.  <!-- nyphud -->
-   Use index slicing to access larger pieces of a string.  <!--  -->
-   Import a function from a library and use it, such as from `math`.  <!-- syssev -->

Double-click to write **your name and Student ID** here: 李瑞琦 3180111638

### Recap: How to use the online Python notebook server.

The following grey cell, in contrast, is an executable *code cell*. 

Code cells require you to enter Python code (or just read what's already there), and then to execute the cell when you are ready to proceed with the lab.  To do so, select inside of the cell and press Ctrl+Enter or Shift+Enter.

```python
1 + 2
```

Output:
```
3
```

Note that two things happened.  One is that the `In []` gained a number—this is basically the number of inputs you have executed in the notebook.  The absolute number almost never matters, but the *order* often does.  You can execute cells out of order, so you have to be sure that if you change one thing later cells are in fact written with the updated value in mind.

Alright, now for the good stuff!

### How numbers (and other values) work

#### Numbers


In the previous lab, although you are not told, you have already encountered integers and floating-point numbers, the basic numeric data types of Python. Let us look at these and more data types, which are simply kinds of data you can describe using Python. You can determine the type of a number by using the intuitively-named type function:

```python
type(2)
```

Output:
```
int
```

```python
type(2.)    # Again, note that `2` is an `int` while `2.` is a `float`.
```

Output:
```
float
```

```python
type(-1)
```

Output:
```
int
```

```python
type(1.564e15)
```

Output:
```
float
```

No surprises there.  *Promotion* is the phenomenon which occurs when (compatible) data types of different kinds are operated on together.  For instance, a `float` plus an `int` is a `float`:

```python
type(1e5+1)
```

Output:
```
float
```

Sometimes functions change a number type, even if the result could be represented as the original type (`int`).  This can occur when intermediate steps in a calculation use `float`.  Sometimes they don't.

```python
from math import exp  # include a library function for the exponential
print('exp(0) = ', exp(0))
type(exp(0))
```

Output:
```
exp(0) =  1.0
float
```

```python
type(abs(-1))
```

Output:
```
int
```

```python
type(abs(-1.))
```

Output:
```
float
```

<!-- grade_id: parameterize-test, points: 1, grade: true, solution: false -->
### Question 1 [5 points]
What is the type of the result of the following expression?
        
        int(exp(abs(-2)) + abs(-1.5)) // 4.5

```python
import math # store the result in the variable T and find T's data type
T = int ( math.exp ( abs(-2) ) + abs(-1.5) ) // 4.5 # your code here
type(T)

```

Output:
```
float
```

<!-- grade_id: parameterize-test, points: 1, grade: true, solution: false -->
### Question 2 [5 points]
What is the type of the result of the following expression?
        
        abs(int(1.5e4))

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
# store the result in the variable S and find S's data type
S = abs ( int(1.5e4) )# your code here
type(S)
```

Output:
```
int
```

### How to write text

Numbers are never enough—you use them to get the math right, but when it comes time to communicating the results to the outside world, you need text.  In programming terms, we refer to a block of text as a *string* (from a _string of characters_).  Some messages are hard-coded, like error messages; others will be generated as necessary by the program.

Since we communicate with the Python intepreter using text commands, we also need a way to distinguish between instructions, the machine (code), and messages written as text (strings).  In other words, we need Python to distinguish between

```python
print(hello world) which will produce a syntax error
```

Output:
```
File "<ipython-input-18-ce2e317b0c8b>", line 1
    print(hello world) which will produce a syntax error
                    ^
SyntaxError: invalid syntax
```

and

```python
print('hello world')
```

Output:
```
hello world
```

The string holds the words within quotes as a single unit.  We typically write strings by surrounding our text in single quotes `''` or double quotes `""`:

```python
'All that glitters is not gold. Often have you heard that told. Many a man his life hath sold. But my outside to behold.'
```

Output:
```
'All that glitters is not gold. Often have you heard that told. Many a man his life hath sold. But my outside to behold.'
```

```python
'Who knows what evil lurks in the hearts of men?  The Shadow knows.'
```

Output:
```
'Who knows what evil lurks in the hearts of men?  The Shadow knows.'
```

```python
'I'm sorry, CS101, I'm afraid I can't do that.'
```

Output:
```
File "<ipython-input-22-9bd7e4ee2112>", line 1
    'I'm sorry, CS101, I'm afraid I can't do that.'
       ^
SyntaxError: invalid syntax
```

That last one has a problem, doesn't it? We wrote a string by marking the ends with single quotes ', but then used a single quote inside of the string.  This thoroughly confused Python, as you can see by the red and black text mixed.

This is why Python lets you use either single or double quotes:

```python
"I'm in love with CS101."
```

Output:
```
"I'm in love with CS101."
```

```python
'You must be "Crazy"'
```

Output:
```
'You must be "Crazy"'
```

Write the following phrases as Python strings.  (You can copy and paste the text itself into the string.)

<!-- grade_id: parameterize-test, points: 1, grade: true, solution: false -->
### Question 3 [5 points]
William Shakespeare's play - The Merchant of Venice 

```python
"William Shakespeare's play - The Merchant of Venice"#Copy and paste using correct quotes

```

Output:
```
"William Shakespeare's play - The Merchant of Venice"
```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 4 [5 points]
- This is a straight double quote, ", while these are curly double quotes: “, ”.

<!-- grade_id: string2, points: 1, grade: true, solution: true -->
```python
'''This is a straight double quote, ", while these are curly double quotes: “, ”.'''#Copy and paste using correct quotes

```

Output:
```
'This is a straight double quote, ", while these are curly double quotes: “, ”.'
```

If things get really hairy, you can also use triple-quotes to mark a string:

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 5 [5 points]
- "Are we finishing yet?" asked an impatient ZJUI's student

```python
'''"Are we finishing yet?" asked an impatient ZJUI's student'''#Copy and paste using correct quotes

```

Output:
```
'"Are we finishing yet?" asked an impatient ZJUI\'s student'
```

<div class="alert alert-info">
Notice that the string returned by Python consists of *single quotes* wrapping a sentence containing *double quotes*, which then contain more *single quotes*.  The inner single quotes have a backslash in front of them—this is called an *escape character*, which is a fancy way of saying that the backslash tells Python to treat the single quote as a 'quote' rather than as marking the beginning or end of a string.
</div>

### How to store a value for future use

In physics and engineering, we often worked with labeled unknowns, called *variables*.  This lets us refer to physical or mathematical values without knowing the value ahead of time.  For instance, the ideal gas law represents the relationship between pressure $P$, temperature $T$, and volume $V$, written as

$$
PV \propto T \text{.}
$$

We can express the relationship between these physical quantities without referring to specific values.

For a computer, variables work much the same way—you can use them to write expressions and formulas, and you can use them to store the results of past calculations.  You create a variable using the assignment operator `=` and reuse it using the same name in subsequent expressions:

```python
x = 5
```

```python
x
```

Output:
```
5
```

```python
x**2
```

Output:
```
25
```

```python
x%2
```

Output:
```
1
```

You update the value of a variable by simply overwriting it:

```python
x = 4.5
```

```python
x
```

Output:
```
4.5
```

```python
x**2
```

Output:
```
20.25
```

In fact, you can even use a variable's *former* value to give it a *new* value:

```python
x = 2*x
```

If you missed it in the first Python lecture, stop and think about this for a moment, because it's very important.  If you were to write down the mathematical statement

$$
x = 2x \text{,}
$$

there is *no nontrival (nonzero) solution*—$x$ divides out of both sides and you are left with the mathematical absurdity $1=2$.  You should *not think* of the assignment operator as reading "equals".  That is, `x = 2x` should be read by you as "`x` is made equal to the value of 2 times `x`".

<div class="alert alert-warning">
The *assignment* operator is a single equals sign, `=`, read "is made equal to".
<br/>
The *equality* operator is a double equals sign, `==`, read "equals".
</div>

As you saw in the reading, you can use much more involved variable names, which is a great idea if you don't want to lose track of what you're doing.

```python
trees_observed = 4
max_torque = 5.0
middleInitial = 'L'
finalSum = 107
way_to_Tipperary = 'east-north-east'
gretting_msg = "Call me Ishmael."
```

A few possible names do not, in fact, work (The following should throw a `Syntax Error`.  We will examine error types later in this course) :

```python
4th = 16.5
```

Output:
```
File "<ipython-input-40-e00999a6b825>", line 1
    4th = 16.5
      ^
SyntaxError: invalid syntax
```

```python
$ = 2
```

Output:
```
File "<ipython-input-41-0dd7977ecaa7>", line 1
    $ = 2
    ^
SyntaxError: invalid syntax
```

<div class="alert alert-warning">
The error that occurs in both of these cases, `SyntaxError`, generally means that you've mistyped something or tried to express something impossible by the rules of Pythonic syntax. Basically, as long as you use only letters, numbers, and the underscore `_` in your variable names, you'll be fine.
</div>

### How variables make your life easier

Imagine writing a long script which depends on the value of a heat transfer coefficient with the value $7.9 \,\text{watt}\cdot\text{meter}^{-2}\cdot\text{kelvin}^{-1}$.  You write a statement to calculate the heat transfer through the wall with a difference of $\Delta T = 1.2 \,\text{kelvin}$ using the value `7.9`:

```python
print("The heat flux is",(7.9 * 1.2),"W/m**2." )
```

Output:
```
The heat flux is 9.48 W/m**2.
```

You may have already noticed the problem here—what happens when the problem changes slightly and you need to use a different heat transfer coefficient or a different $\Delta T$?  This is where variables shine:

$$
q = h \Delta T
$$

If you set things up properly, you will only have to define a value for $h$ and $\Delta T$ once, at the beginning of your script, and it will be the same throughout the program.  Otherwise, you might have had to search and replace every single instance—and hope you didn't miss one!

```python
h = 7.9     # heat transfer coef, W/m^2/K
DT = 1.2    # difference in temperature between fluid and surface, K
q = h * DT  # heat flux or amount of heat transferred, W/m^2

print("The heat transfer is",q,"W/m**2.")

DT = 1.3
q = h * DT  # heat flux or amount of heat transferred, W/m^2

print("The heat transfer is",q,"W/m**2.")
```

Output:
```
The heat transfer is 9.48 W/m**2.
The heat transfer is 10.270000000000001 W/m**2.
```

<div class="alert alert-info">
Note as well that since physical quantities always have physical dimensions, you should always document these with the value (or include a note to the effect that all values are [SI](https://en.wikipedia.org/wiki/International_System_of_Units), [CGS](https://en.wikipedia.org/wiki/Centimetre%E2%80%93gram%E2%80%93second_system_of_units), [FPS](https://en.wikipedia.org/wiki/Foot%E2%80%93pound%E2%80%93second_system), etc.).
<pre><code>max_torque = 5.0  # N*m</code></pre>
A brief explanation may be helpful as well:
<pre><code>dx = 1e-2  # mesh discretization, m</code></pre>
Form this habit now and it will serve you well as you begin writing code, scripts, and longer programs.
</div>

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 6 [20 points]
-   Correct the following code to use variables `v` and `a` instead of hard-coded values.

<!-- grade_id: parameterize, points: 0, grade: false, solution: true -->
```python
# modify this code
# the base equation is:
# d = v*t + 0.5*a*t*t
v = 5.0
a = -9.8
t = 10.0

# Correct below to include the variables:
d = v*t + 0.5*a*t**2
```

<!-- grade_id: parameterize-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(d, -440)
print('Success!')
```

Output:
```
Success!
```

#### Strings & Indexing

The basic text data type, the `string`, has already shown up in a few places:
    
    print("I was sleeping through this morning's lecture")

In order to distinguish the human-readable message from the machine command `print`, we surrounded the string with some set of quotes.  *Indexing* means selecting a single letter in a string (or more generally, an element of a sequence).  We do this by placing a bracketed number after a string variable:

```python
message = "They were sleeping through this morning's lecture."
print(message[3])
```

Output:
```
y
```

Recall why `message[3]` returns `'y'` not `'e'`.  Typically we would consider `'y'` to be the fourth element of the string, but many programming languages—including Python—instead start counting at zero. In this model, the string message is indexed as follows:

| index  | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | ... | 46 | 47 | 48 | 49 |
|--------|---|---|---|---|---|---|---|---|---|-----|----|----|----|----|
| letter | T | h | e | y |   | w | e | r | e | ... | u  | r  | e  | .  |

Aside from starting with zero, indexing is pretty straightforward.

<div class="alert alert-info">
<p>
Programming languages that count from zero are referred to as *zero-indexed*, while those that start from one are *one-indexed*.  The rationale for zero-indexing is to treat the first letter as the beginning of the string. Thus `'h'` is *+1* away from the beginning, `'w'` is *+2* away, etc.  Languages that are “related” to [C](https://en.wikipedia.org/wiki/C_(programming_language) generally adopt this convention.
</p><p>
One argument for one-indexing is that humans (and mathematicians) count from one "naturally".  Languages related to [Fortran](https://en.wikipedia.org/wiki/Fortran) often use one-indexing.
</p><p>
If you'd like to more information on this, please refer to [this Q&A at Stack Overflow](https://programmers.stackexchange.com/questions/110804/why-are-zero-based-arrays-the-norm).
</p>
</div>

```python
message[10]
```

Output:
```
's'
```

```python
message[60]
```

Output:
```
IndexError                                Traceback (most recent call last)
<ipython-input-52-1d4f2b183760> in <module>()
----> 1 message[60]

IndexError: string index out of range
```

(In the last case, we tried to index past the end of the string—we can't do that!)

Regular indexing is a fairly natural process for us: it's just counting off with whole numbers. What happens if you use a negative index?

```python
message[-1]
```

Output:
```
'.'
```

Negative indexing counts down from the end of the string. This is very convenient and saves you the trouble of figuring out the length of the string and then subtracting from that.

Some things probe the limits of sanity, though:

```python
message[8.5]
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-54-f75870a787a2> in <module>()
----> 1 message[8.5]

TypeError: string indices must be integers
```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 7 [5 points]
What is the index of the following letters in this string?

"You're not fully clean till you're Zestfully clean."
- `'n'` (first occurance of n)

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
s = "You're not fully clean till you're Zestfully clean."# your answer here as the variable ans,
ans = s [7]
```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 8 [5 points]
- `'y'` (last ocuurance — practise your negative indexing)

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
# your answer here as a variable ans
ans = s[-8]
```

#### Slicing

We can generalize our notion of *indexing* to *slicing*.  Slicing means using a *range of indices* to take a *slice* of a string.  A range is written `string[left:right]`, where the `left` index is part of the string but the `right` index is not.  (We call this a *half-open* range.)

Remember that Python starts counting from zero, and that while the left-hand side is included, the right-hand side is _not_.

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
pangram = 'Sphinx of black quartz: judge my vow.'
```

```python
pangram[0:5]
```

Output:
```
'Sphin'
```

```python
pangram[6:13]
```

Output:
```
' of bla'
```

```python
pangram[0:0]  # Note that the right-hand index is not in the range, so this string is empty.
```

Output:
```
''
```

```python
pangram[0:-1]
```

Output:
```
'Sphinx of black quartz: judge my vow'
```

```python
pangram[5:]  # We don't have to specify both sides; if we don't, Python assumes we are going all the way to the end or coming all the way from the beginning.
```

Output:
```
'x of black quartz: judge my vow.'
```

```python
pangram[:10]
```

Output:
```
'Sphinx of '
```

```python
pangram[:]  # Thus this is just a copy of the entire original string.
```

Output:
```
'Sphinx of black quartz: judge my vow.'
```

```python
pangram[::2] # Take every other character
```

Output:
```
'Shn fbakqat:jdem o.'
```

Although strings *are* important and useful, you're still fairly limited in what you're able to do with them.  But many aspects of indexing will carry over to data types you'll meet later, like `list`s and `ndarray`s ($n$-dimensional arrays)—you will use these techniques to manipulate large arrays of numerical data as well.

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 9 [5 points]
-   How would you extract the letters `'ssupdla nioiaim'` from the following string?
        
        showoff = 'sesquipedalian antinominalism'

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
# definition
showoff = 'sesquipedalian antinominalism'
```

<!-- grade_id: slice, points: 1, grade: true, solution: true -->
```python
# your answer here
showoff[ ::2 ]
```

Output:
```
'ssupdla nioiaim'
```

#### `math` Library Functions

*Modules* are collections of functions which allow you to take advantage of the collective work of the scientific and engineering community when carrying out your calculations or building your models.  One of the strengths of Python is its enormous and rich library of modules, which provide both basic and involved operations far beyond the built-in operators.

Let's consider the `math` library.  To use functions like `sin`, `cos`, `exp`, and `log`, you simply need to `import math`:

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
import math
```

You still need to specify *which* module contains the functions to use them, but otherwise it's quite straightforward:

$$
\sin\left(\frac{\pi}{4}\right)
$$

```python
math.sin(math.pi/4)
```

Output:
```
0.7071067811865476
```

$$
\left[\text{ceil}\left(4.5\right)\right] !
$$

```python
math.factorial(math.ceil(4.5))
```

Output:
```
120
```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 10 [10 points]
-   Write an expression for the following formula:
    
    $$
    \frac{(m+n)!}{(m+1)(n-1)!}
    $$
    
    (Note that the factorial refers only to the second term of the denominator, not the entire denominator.)

```python
# these are representative values—don't change these
m = 2
n = 3
```

<!-- grade_id: form, points: 0, grade: false, solution: true -->
```python
# write your expression here
expr = math.factorial (m + n) / ( (m + 1) * math.factorial(n - 1) )# your code here
```

<!-- grade_id: form-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(expr, 20.0)
print('Success!')
```

Output:
```
Success!
```

Some modules are quite large, and you may not want to import everything into memory at once.  In that case, you can import a single function from a module as follows:

```python
from math import asin  # arcsine
```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
### Question 11 [5 points]
-   What line of Python code could you write to import `log10` from `math`?  (This is the logarithm base 10, $\log_{10}(x)$.)

<!-- grade_id: import, points: 0, grade: false, solution: true -->
```python
from math import log10# write your expression here

```

<!-- grade_id: import-test, points: 1, grade: true, solution: false -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
try:
    log10(0.25)
except NameError:
    raise AssertionError('log10 not imported yet')
print('Success!')
```

Output:
```
Success!
```

You've seen basic expressions in the lab and in the homework.  Let's build some more complicated expressions now to serve us in real model-building later.

### Question 12 [15 points]
-   Write an expression describing the period of a pendulum for small amplitudes,
    $$
    T = 2\pi \sqrt\frac{L}{g} \text{.}
    $$

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
from math import pi
L = 0.72  # m
g = 9.8   # acceleration due to gravity.
```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
T = 2 * pi * ( L / g ) ** 0.5# write a formula for T here, using pi from the math library
# you can also use "from math import sqrt" if you like, or use **0.5 as an equivalent

```

<!-- grade_id: string1, points: 1, grade: true, solution: true -->
```python
# it should pass this test---do NOT edit this cell
from numpy import isclose
assert isclose(T, 1.70307)
print('Success!')
```

Output:
```
Success!
```

# Before you submit...

### Submission [10 points]
Make sure that you have filled in any code block which is used for answering a Question #.

Save this file as lab04-emailID-studentID.ipynb then UPLOAD and EMAIL to your TAs!